---
title: Option Market Mechanics and Basic Trading Strategies
description: >-
  This chapter provides an overview of option market mechanics, such as the
  different type of options, margin accounting, the options clearing house, and
  the specification of stock options contracts. It will also cover the different
  type of options positions and how they are used for trading.

--- type:VideoExercise lang:r xp:50 skills:1 key:687d3c7fa8
## Types of Options and Option Positions



*** =video_link
//player.vimeo.com/video/154783078

*** =video_hls
//videos.datacamp.com/transcoded/000_placeholders/v1/hls-temp.master.m3u8

--- type:MultipleChoiceExercise lang:r xp:50 skills:1 key:d8e6ca9307
## The Difference European and American Options

Which of the following pairs of characteristics distringuish European and American options?

*** =instructions

*** =hint

*** =pre_exercise_code
```{r}

```

*** =sct
```{r}

```

--- type:NormalExercise lang:r xp:100 skills:1 key:c88ac04897
## Visualizing Options Positions

In this exercise you will plot the payoff functions for different options positions. rep(), plot(), seq()

*** =instructions

*** =hint

*** =pre_exercise_code
```{r}

```

*** =sample_code
```{r}

```

*** =solution
```{r}

```

*** =sct
```{r}

```

--- type:VideoExercise lang:r xp:50 skills:1 key:7da4719903
## Commissions, Margin Requirements and the Clearing Corporation

Discuss some of the crucial characteristics of option markets institutions

*** =video_link
//player.vimeo.com/video/154783078

*** =video_hls
//videos.datacamp.com/transcoded/000_placeholders/v1/hls-temp.master.m3u8

--- type:MultipleChoiceExercise lang:r xp:50 skills:1 key:ca885ee5bb
## The Option Clearing Corporation

Which of the following is correct concerning the Options Clearing Corporation?

*** =instructions

*** =hint

*** =pre_exercise_code
```{r}

```

*** =sct
```{r}

```

--- type:MultipleChoiceExercise lang:r xp:50 skills:1 key:5897c06d26
## Understanding Margin Requirements

Which of the following explains why margin accounts are necessary for traders when they write options but not when they buy options?

*** =instructions

*** =hint

*** =pre_exercise_code
```{r}

```

*** =sct
```{r}

```

--- type:VideoExercise lang:r xp:50 skills:1 key:61210b77d7
## Basic Option Trading Strategies

Discuss covered calls, protective puts, and synthetic calls, puts, and forwards

*** =video_link
//player.vimeo.com/video/154783078

*** =video_hls
//videos.datacamp.com/transcoded/000_placeholders/v1/hls-temp.master.m3u8

--- type:MultipleChoiceExercise lang:r xp:50 skills:1 key:c187abac7f
## Understanding Covered Calls and Protective Puts

In the example above, how much profit would David have made had he not written the covered call and Google stock went up to $90?

*** =instructions

*** =hint

*** =pre_exercise_code
```{r}

```

*** =sct
```{r}

```

--- type:NormalExercise lang:r xp:100 skills:1 key:b035053535
## Plotting Covered Calls and Protective Puts

In this exercise you will write R functions to plot basic option trading strategies. seq(), rep(), plot()

*** =instructions

*** =hint

*** =pre_exercise_code
```{r}

```

*** =sample_code
```{r}

```

*** =solution
```{r}

```

*** =sct
```{r}

```

--- type:VideoExercise lang:r xp:50 skills:1 key:4d7e6d3a04
## More Advanced Option Trading Strategies

Straddles, Strangles, and Spreads

*** =video_link
//player.vimeo.com/video/154783078

*** =video_hls
//videos.datacamp.com/transcoded/000_placeholders/v1/hls-temp.master.m3u8

--- type:MultipleChoiceExercise lang:r xp:50 skills:1 key:796859ee98
## Understanding Spread Trades

Consider the position of selling (instead of buying) a bull spread - that is, writing a call with a low strike price and buying a call with a high strike price. Which of the following explains how this differs from buying a bull spread?

*** =instructions

*** =hint

*** =pre_exercise_code
```{r}

```

*** =sct
```{r}

```

--- type:NormalExercise lang:r xp:100 skills:1 key:8e98873020
## Plotting Spread Trades

In this exercise you will write an R function to plot basic option spread trades. rep(), seq(), plot()

*** =instructions

*** =hint

*** =pre_exercise_code
```{r}

```

*** =sample_code
```{r}

```

*** =solution
```{r}

```

*** =sct
```{r}

```
