---
title: Put-Call Parity
description: >-
  This chapter provides an overview of the crucial concept of put-call parity
  for European options

--- type:VideoExercise lang:r xp:50 skills:1 key:ce1eae6a4a
## Put-Call Parity

Explain the concept of put-call parity

*** =video_link
//player.vimeo.com/video/154783078

*** =video_hls
//videos.datacamp.com/transcoded/000_placeholders/v1/hls-temp.master.m3u8

--- type:MultipleChoiceExercise lang:r xp:50 skills:1 key:12e09d1896
## Understanding Parity Relations

Suppose that the asset is at $1020 and the interest rate is 5%. At what strike price should a one-year call price equal the one-year put price?

*** =instructions

*** =hint

*** =pre_exercise_code
```{r}

```

*** =sct
```{r}

```

--- type:NormalExercise lang:r xp:100 skills:1 key:a7af7c70bf
## Implementing Put-Call Parity

In this exercise you will write an R function to implement Put-Call parity. rep(), seq()

*** =instructions

*** =hint

*** =pre_exercise_code
```{r}

```

*** =sample_code
```{r}

```

*** =solution
```{r}

```

*** =sct
```{r}

```

--- type:VideoExercise lang:r xp:50 skills:1 key:6575a137a6
## Generalized Parity Relationships

Explain the concept of generalized parity relations

*** =video_link
//player.vimeo.com/video/154783078

*** =video_hls
//videos.datacamp.com/transcoded/000_placeholders/v1/hls-temp.master.m3u8

--- type:MultipleChoiceExercise lang:r xp:50 skills:1 key:006b00cecd
## Understanding Generalized Parity Relationships

Suppose call and put prices are given by (show). Which of the following no-arbitrage properties is violated?

*** =instructions

*** =hint

*** =pre_exercise_code
```{r}

```

*** =sct
```{r}

```

--- type:NormalExercise lang:r xp:100 skills:1 key:688f4d9116
## Implementing Generalized Parity Relationships

In this exercise you write an R function to test for generalized parity relations. rep(), seq()

*** =instructions

*** =hint

*** =pre_exercise_code
```{r}

```

*** =sample_code
```{r}

```

*** =solution
```{r}

```

*** =sct
```{r}

```

--- type:VideoExercise lang:r xp:50 skills:1 key:231a60dbf7
## Comparing Options with Respect to Style, Maturity, and Strike

Explain the variation of option contracts with respect to key features of options

*** =video_link
//player.vimeo.com/video/154783078

*** =video_hls
//videos.datacamp.com/transcoded/000_placeholders/v1/hls-temp.master.m3u8

--- type:MultipleChoiceExercise lang:r xp:50 skills:1 key:5c5da01f3b
## Understanding Key Features of Options

Which of the following is true concerning option style, maturity and strike price?

*** =instructions

*** =hint

*** =pre_exercise_code
```{r}

```

*** =sct
```{r}

```

--- type:NormalExercise lang:r xp:100 skills:1 key:ace21d285b
## Testing the Implications of Differences in Maturity and Strike

In this exercise you will write R functions to test for the difference in option values with respect to maturity and strike. rep(), seq()

*** =instructions

*** =hint

*** =pre_exercise_code
```{r}

```

*** =sample_code
```{r}

```

*** =solution
```{r}

```

*** =sct
```{r}

```

--- type:VideoExercise lang:r xp:50 skills:1 key:d994c7c259
## Parity Bounds for American Options

Explain parity bounds for American options

*** =video_link
//player.vimeo.com/video/154783078

*** =video_hls
//videos.datacamp.com/transcoded/000_placeholders/v1/hls-temp.master.m3u8

--- type:MultipleChoiceExercise lang:r xp:50 skills:1 key:ca643194c3
## Understanding American Parity Bounds

Which of following is correct concerning American parity bounds?

*** =instructions

*** =hint

*** =pre_exercise_code
```{r}

```

*** =sct
```{r}

```
