---
title: Introduction to Derivatives Markets
description: >-
  This chapter provides an introduction to option contracts, option markets, and
  the foundational economic principle of the law of one price otherwise known as
  the no-arbitrage principle.

--- type:VideoExercise lang:r xp:50 skills:1 key:647d523c0d
## An Introduction to Financial Derivatives

Introduce financial derivative contracts with special emphasis on options

*** =video_link
//player.vimeo.com/video/154783078

*** =video_hls
//videos.datacamp.com/transcoded/000_placeholders/v1/hls-temp.master.m3u8

--- type:MultipleChoiceExercise lang:r xp:50 skills:1 key:f1f10d197c
## The Unique Feature of Options Contracts

Which of the following is the primary distinguishing characteristic of options contracts

*** =instructions

*** =hint

*** =pre_exercise_code
```{r}

```

*** =sct
```{r}

```

--- type:VideoExercise lang:r xp:50 skills:1 key:cf2a74bf4b
## The Social Function of Financial Deratives

Discuss the social benefits of financial derivatives

*** =video_link
//player.vimeo.com/video/154783078

*** =video_hls
//videos.datacamp.com/transcoded/000_placeholders/v1/hls-temp.master.m3u8

--- type:MultipleChoiceExercise lang:r xp:50 skills:1 key:baec93ca48
## The Social Benefits of Derivatives Markets

Which of the following is an example of the social benefit of derivatives markets?

*** =instructions

*** =hint

*** =pre_exercise_code
```{r}

```

*** =sct
```{r}

```

--- type:MultipleChoiceExercise lang:r xp:50 skills:1 key:564e0a1226
## The Social Benefits of Derivatives Markets

Due to a flood in 1991 the Chicago Board of Trade stoped trading. During this time grain elevators refused to post spot prices. Which of the following properties of derivatives markets explains this?

*** =instructions

*** =hint

*** =pre_exercise_code
```{r}

```

*** =sct
```{r}

```

--- type:MultipleChoiceExercise lang:r xp:50 skills:1 key:0a878f55b9
## The Informational Role of Options Prices

Which of the following demonstrates how option prices convey information about the economy more boardly?

*** =instructions

*** =hint

*** =pre_exercise_code
```{r}

```

*** =sct
```{r}

```

--- type:VideoExercise lang:r xp:50 skills:1 key:39b9ee3510
## Call and Put Option Payoffs

Explain the payoff functions of vanilla call and put options and how they differ from forward payoffs

*** =video_link
//player.vimeo.com/video/154783078

*** =video_hls
//videos.datacamp.com/transcoded/000_placeholders/v1/hls-temp.master.m3u8

--- type:MultipleChoiceExercise lang:r xp:50 skills:1 key:b5d74c1b97
## Call vs Put Payoff

Which of the following pairs of traits differentiates call options from put payoffs

*** =instructions

*** =hint

*** =pre_exercise_code
```{r}

```

*** =sct
```{r}

```

--- type:NormalExercise lang:r xp:100 skills:1 key:869de598e5
## Write Call and Put Option Payoff Function

In this exercise you will write R functions to calculate the payoff for a plain vanilla call and put options. max(), min(), pmax(), pmin()

*** =instructions

*** =hint

*** =pre_exercise_code
```{r}

```

*** =sample_code
```{r}

```

*** =solution
```{r}

```

*** =sct
```{r}

```

--- type:VideoExercise lang:r xp:50 skills:1 key:2f0e5689e6
## Visualizing Call and Put Payoff Functions

Discuss the structure of call and put payoff functions and how we can visualize them

*** =video_link
//player.vimeo.com/video/154783078

*** =video_hls
//videos.datacamp.com/transcoded/000_placeholders/v1/hls-temp.master.m3u8

--- type:NormalExercise lang:r xp:100 skills:1 key:4c00383a07
## Plot Call and Put Payoff Functions

In this exercise you will plot call and put payoff functions to help visualize their respective payoffs. plot(), seq(), rep()

*** =instructions

*** =hint

*** =pre_exercise_code
```{r}

```

*** =sample_code
```{r}

```

*** =solution
```{r}

```

*** =sct
```{r}

```
