---
title: The Black-Scholes Model
description: >-
  This chapter provides an introduction to the Black-Scholes Options Pricing
  Model.

--- type:VideoExercise lang:r xp:50 skills:1 key:01cd460d39
## Introduction to the Black-Scholes Formula for Call and Put Options

Present the Black-Scholes formula for European call and puts options, and explain its use for pricing and trading options

*** =video_link
//player.vimeo.com/video/154783078

*** =video_hls
//videos.datacamp.com/transcoded/000_placeholders/v1/hls-temp.master.m3u8

--- type:MultipleChoiceExercise lang:r xp:50 skills:1 key:6b48fbaf38
## Comparing the Black-Scholes and Binomial Option Pricing Models

Which pair of statements is true concerning the Black-Scholes and Binomial Options Pricing Models?

*** =instructions

*** =hint

*** =pre_exercise_code
```{r}

```

*** =sct
```{r}

```

--- type:NormalExercise lang:r xp:100 skills:1 key:317ef67097
## Implementing Black-Scholes

In this exercise you will write R functions to implement the Black-Scholes model for call and put options. log(), exp(), pnorm()

*** =instructions

*** =hint

*** =pre_exercise_code
```{r}

```

*** =sample_code
```{r}

```

*** =solution
```{r}

```

*** =sct
```{r}

```

--- type:VideoExercise lang:r xp:50 skills:1 key:776bcd4f52
## The Option Greeks

Present the so-called option greek parameters and explain how they measure risk sensitivities for options

*** =video_link
//player.vimeo.com/video/154783078

*** =video_hls
//videos.datacamp.com/transcoded/000_placeholders/v1/hls-temp.master.m3u8

--- type:MultipleChoiceExercise lang:r xp:50 skills:1 key:73360c4050
## Understanding The Greek Parameters

Which of the following pairs of statements is correct concerning the greek parameters?

*** =instructions

*** =hint

*** =pre_exercise_code
```{r}

```

*** =sct
```{r}

```

--- type:VideoExercise lang:r xp:50 skills:1 key:35fc78a04f
## Implied Volatility

Discuss the calculation, interpretation, and use of implied volatities from the Black-Scholes model

*** =video_link
//player.vimeo.com/video/154783078

*** =video_hls
//videos.datacamp.com/transcoded/000_placeholders/v1/hls-temp.master.m3u8

--- type:MultipleChoiceExercise lang:r xp:50 skills:1 key:ca5d9cde26
## Understanding Implied Volatility

Which of the following statements correctly characterizes implied volatlity?

*** =instructions

*** =hint

*** =pre_exercise_code
```{r}

```

*** =sct
```{r}

```

--- type:NormalExercise lang:r xp:100 skills:1 key:8a3fc23e68
## Implementing Implied Volatility Calculation

In this exercise you will write and R function to calculate the Black-Scholes implied volatlity using the binary search method. log(), exp(), pnorm(), for loops

*** =instructions

*** =hint

*** =pre_exercise_code
```{r}

```

*** =sample_code
```{r}

```

*** =solution
```{r}

```

*** =sct
```{r}

```

--- type:VideoExercise lang:r xp:50 skills:1 key:004572e140
## Using Implied Volatility in Trading



*** =video_link
//player.vimeo.com/video/154783078

*** =video_hls
//videos.datacamp.com/transcoded/000_placeholders/v1/hls-temp.master.m3u8

--- type:MultipleChoiceExercise lang:r xp:50 skills:1 key:3d34333909
## Understanding Implied Volatility for Trading

Why do option traders prefer to use implied volatlities over prices when trading options?

*** =instructions

*** =hint

*** =pre_exercise_code
```{r}

```

*** =sct
```{r}

```

--- type:NormalExercise lang:r xp:100 skills:1 key:eaad100c43
## Simulated Trading with Implied Volatility

In this exercise you will write R code to use implied volatility in a trading simulation. log(), exp(), rnorm(), for loops

*** =instructions

*** =hint

*** =pre_exercise_code
```{r}

```

*** =sample_code
```{r}

```

*** =solution
```{r}

```

*** =sct
```{r}

```
