---
title: The Black-Scholes Model
description: >-
  This chapter provides an introduction to the Black-Scholes Options Pricing
  Model.

--- type:VideoExercise lang:r xp:50 skills:1 key:43e3aafd64
## Introduction to the Black-Scholes Formula for Call and Put Options

Present the Black-Scholes formula for European call and puts options, and explain its use for pricing and trading options

*** =video_link
//player.vimeo.com/video/154783078

*** =video_hls
//videos.datacamp.com/transcoded/000_placeholders/v1/hls-temp.master.m3u8

--- type:MultipleChoiceExercise lang:r xp:50 skills:1 key:4f78516777
## Comparing the Black-Scholes and Binomial Option Pricing Models

Which pair of statements is true concerning the Black-Scholes and Binomial Options Pricing Models?

*** =instructions

*** =hint

*** =pre_exercise_code
```{r}

```

*** =sct
```{r}

```

--- type:NormalExercise lang:r xp:100 skills:1 key:ea6da51287
## Implementing Black-Scholes

In this exercise you will write R functions to implement the Black-Scholes model for call and put options. log(), exp(), pnorm()

*** =instructions

*** =hint

*** =pre_exercise_code
```{r}

```

*** =sample_code
```{r}

```

*** =solution
```{r}

```

*** =sct
```{r}

```

--- type:VideoExercise lang:r xp:50 skills:1 key:0cefc9659e
## The Option Greeks

Present the so-called option greek parameters and explain how they measure risk sensitivities for options

*** =video_link
//player.vimeo.com/video/154783078

*** =video_hls
//videos.datacamp.com/transcoded/000_placeholders/v1/hls-temp.master.m3u8

--- type:MultipleChoiceExercise lang:r xp:50 skills:1 key:927897f9b6
## Understanding The Greek Parameters

Which of the following pairs of statements is correct concerning the greek parameters?

*** =instructions

*** =hint

*** =pre_exercise_code
```{r}

```

*** =sct
```{r}

```

--- type:VideoExercise lang:r xp:50 skills:1 key:df50595cbb
## Implied Volatility

Discuss the calculation, interpretation, and use of implied volatities from the Black-Scholes model

*** =video_link
//player.vimeo.com/video/154783078

*** =video_hls
//videos.datacamp.com/transcoded/000_placeholders/v1/hls-temp.master.m3u8

--- type:MultipleChoiceExercise lang:r xp:50 skills:1 key:ca194f38a3
## Understanding Implied Volatility

Which of the following statements correctly characterizes implied volatlity?

*** =instructions

*** =hint

*** =pre_exercise_code
```{r}

```

*** =sct
```{r}

```

--- type:NormalExercise lang:r xp:100 skills:1 key:6fa8cbfade
## Implementing Implied Volatility Calculation

In this exercise you will write and R function to calculate the Black-Scholes implied volatlity using the binary search method. log(), exp(), pnorm(), for loops

*** =instructions

*** =hint

*** =pre_exercise_code
```{r}

```

*** =sample_code
```{r}

```

*** =solution
```{r}

```

*** =sct
```{r}

```

--- type:VideoExercise lang:r xp:50 skills:1 key:324bc537b7
## Using Implied Volatility in Trading



*** =video_link
//player.vimeo.com/video/154783078

*** =video_hls
//videos.datacamp.com/transcoded/000_placeholders/v1/hls-temp.master.m3u8

--- type:MultipleChoiceExercise lang:r xp:50 skills:1 key:a76033660b
## Understanding Implied Volatility for Trading

Why do option traders prefer to use implied volatlities over prices when trading options?

*** =instructions

*** =hint

*** =pre_exercise_code
```{r}

```

*** =sct
```{r}

```

--- type:NormalExercise lang:r xp:100 skills:1 key:0c11389eb6
## Simulated Trading with Implied Volatility

In this exercise you will write R code to use implied volatility in a trading simulation. log(), exp(), rnorm(), for loops

*** =instructions

*** =hint

*** =pre_exercise_code
```{r}

```

*** =sample_code
```{r}

```

*** =solution
```{r}

```

*** =sct
```{r}

```
